# *NBGrader*

**Objetivo general**
- Conocer las ventajas de la biblioteca de evaluación *NBGrader* y cómo usarla para evaluar de manera automatizada los conocimientos.

**Objetivos particulares**
- Familiarizarse con la estructura y herramientas de *NBGrader*.
- Generar notebooks que puedan ser evaluados de manera automatizada. 
- Integrar la evaluación automatizada a los materiales generados de manera propia.

## Contenido
- [1 - ¿Qué es *NBGrader* 1.](#1)
    - [1.1 - Estructura de *NBGrader*.](#1-1)
    - [1.2 - Tipo de celdas *NBGrader*.](#1-2)
    - [1.2 - Evaluación automatizada de celdas.](#1-3)
        - [Ejemplo 1 - Evaluación automatizada.](#eje-1)
        - [Ejercicio 1 - Evaluación automatizada *NBGrader*.](#ej-1)
- [2 - Evaluación MACTI.](#2)
    - [2.1 - Estructura evaluación MACTI.](#2-1)
    - [2-2 - Biblioteca *macti_lib*.](#2-2)
    - [2-3 - Método `verifica()`.](#2-3)
        - [Ejercicio 2 - Evaluación automatizada MACTI.](#ej-2)

<a name='1'></a>
## *NBGrader*

*NBGrader* es una biblioteca que nos permite automatizar y organizar de mejor manera la evaluación de los conocimientos generados mediante notebooks.

En general podemos decir que *NBGrader* es un conjunto de herramientas escritas para el lenguaje *Python* que nos permiten automatizar la evaluación de los contenidos. De manera práctica *NBGrader* agrega elementos al proyecto *Jupyter Notebook* para aligerar la tarea de evaluar a los alumnos.

En este <a href="https://youtu.be/5WUm0QuJdFw">link</a> puedes ver un video muy completo sobre el proyecto *NBGrader* y como se usa.

<a name='1-1'></a>
### Estructura *NBGrader*

Una de las primeras cosas de que tenemos que aprender de *NBGrader* es que para evaluar los *notebooks*, vamos a tener 2 versiones del mismo, una de ellas es la versión del profesor (está contiene las preguntas y respuestas) y la versión del estudiante (esta contiene las preguntas sin las respuestas).

Para comenzar con el desarrollo de materiales haciendo uso de *NBGrader* es necesario facilitarse con 3 elementos básicos de la herramienta:

- 1 - *Formgrader*: es un nuevo elemento en la *GUI* de *Jupyter Notebook* que nos permite generar notebooks que serán evaluados de manera automatizada.
- 2 - *source*: es el folder donde se almacena la versión del profesor.
- 3 - *release*: es el folder donde se almacena la versión del estudiante.

<img src="./utils/figs/estructura.png"  hspace="5" vspace="5"/>

Para que se genere el folder *source* (y posteriormente release) es necesario crear un *assignment*, un *assignment* es un *notebook* que queremos que sea evaluado y para crearlo es necesario seleccionar *Formgrader*.

<img src="./utils/figs/formgrader.png"  hspace="5" vspace="5"/>

Posteriormente es necesario seleccionar *Add new asignment*

<img src="./utils/figs/newasignment.png"  hspace="5" vspace="5"/>

Le podemos dar cualquier nombre al *asignment* aunque se recomienda un nombre que facilite identificar de qué se trata el material dentro del *asignment*. En este caso le vamos a dar el nombre *Prueba* y una vez creado el *asignment* el resto es tan sencillo como crear un *notebook* tradicional.

<img src="./utils/figs/nuevonotebook.png"  hspace="5" vspace="5"/>

<a name='1-2'></a>
### Tipo de celdas *NBGrader*

*NBGrader* agrega nuevos tipos de celdas (además de las ya conocidas), mediante las cuales podemos tener más control y sobretodo estas nuevas celdas son las que nos permiten automatizar la evaluación, estas son:

- *readonly*: permite deshabilitar la edición de las celdas, se usa principalmente para celdas de texto.
- *manually graded answer*: son celdas que serán evaluadas de manera manual.
- *manually graded task*: son celdas que serán evaluadas de manera manual y que formarán parte de la evaluación total del curso.
- *auto graded answer*: son celdas que contienen los elementos a evaluar y qué *NBGrader* identifica para eliminar el código necesario.
- *auto graded test*: son celdas que serán evaluadas de manera AUTOMATICA y que le permiten a alumno saber si acerto o tuvo errores y que no pueden ser modificadas por el alumno.

Para tener acceso a estos nuevos tipos de celda, es necesario tener abierto el *notebook* sobre el cual vamos a usar *NBGrader*, seleccionar el menú *View*, después *Cell Toolbar*.

<img src="./utils/figs/evaluacion.gif"  hspace="5" vspace="5"/>

Una vez hecho esto, ya tenemos acceso a los nuevos tipos de celdas previamente descritos.

<a name='1-3'></a>
### Evaluación automatizada de celdas

Ya que tenemos la estructura necesaria y acceso a todas las herramientas de evaluación lo siguiente que necesitamos es identificar el conocimiento que deseamos evaluar y cómo lo vamos a hacer.

Para este ejemplo sencillo vamos a generar una celda de evaluación automatizada (*auto graded answer*) y dejaremos que *NBGrader* se encargue de la evaluación.

**Importante:** *NBGrader* permite automatizar y organizar de manera más completa el proceso de evaluación del conocimiento, pero por motivos de alcance del taller no las veremos. Se recomienda revisar el video mencionado en la parte inicial para más detalle.

Los siguientes pasos para la evaluación los podemos definir de la siguiente manera:

- Identificar el conocimiento a evaluar: puede ser desde un simple cálculo numérico (cálculo de un área, operación aritmética, etc.) hasta un procedimiento más complejo como puede ser un método numérico o un algoritmo más sofisticado.
- Escribir el código asociado a la evaluación: para poder evaluar el conocimiento, es necesario tener la respuesta, por lo que es necesario escribir el código completo (versión del profesor) dentro de una o más celdas a evaluar.
- Seleccionar la sección que debe responder el alumno: partiendo del código escrito en el paso previo, ahora es necesario seleccionar la parte del mismo que va a ser eliminada por *NBGrader* y que debe completar el alumno.
- Contrastar la respuesta correcta contra la del alumno: existen varias formas de llevar a cabo este paso, pero uno de los más sencillos es usar alguna herramienta (por ejemplo `assert`).
- Dejar que *NBGrader* elimine las secciones: ya que sabemos cuales son las secciones a evaluar, *NBGRader* las sustituye por un mensaje de error para que el alumno lo complete.

<a name='eje-1'></a>
#### **Ejemplo 1 - Evaluación automatizada**

Para este ejemplo sencillo, vamos a pedirle al alumno que complete la función `area_cuadrado(lado)`, misma que va a recibir la longitud del lado como parámetro y debe devolver el área del cuadrado.

Una vez que generamos un *notebook* llamado "Evaluacion", dentro del *assignment* "Prueba" dentro del mismo vamos a agregar la (o las) celda a evaluar como se puede ver en la siguiente imagen.

<img src="./utils/figs/correcto.png"  hspace="5" vspace="5"/>

Lo primero que debemos notar en esta imagen (*notebook*) es que la plantilla (estructura de la prueba) ya fue definida por el profesor, es decir que la definición de la función y algunos elementos del código ya fueron definidos y eso permite restringir los posibles errores o confusiones de parte del alumno.

Por otro lado, la sección que debe definir el alumno se encuentra en una celda de tipo *autograded answer* y la celda que se encarga de la evaluación se encuentra en una celda de tipo *autograded test*.

Además es importante notar que la sección que debe completar (y que posteriormente será eliminada por *NBGrader*) debe estar acotada por los comentarios `#BEGIN SOLUTION` y `#END SOLUTION`.

Y lo último que hay que notar es que, en este caso, la implementación (respuesta) es correcta y se muestra un mensaje que le indica al alumno que pasó la prueba.

Sin embargo en la siguiente imágen podemos notar que la respuesta fue incorrecta y se le indica alumno que revise su función.

<img src="./utils/figs/incorrecto.png"  hspace="5" vspace="5"/>

Ya que tenemos la versión del profesor de este *notebook* ("Evaluacion") la cual contiene todos los elementos necesarios y validamos que la prueba funciona de manera correcta, lo siguiente es que *NBGrader* se encargue de generar la versión del estudiante que no contiene las respuestas.

Para hacer esto es necesario volver a la sección *Formgrader* y ubicar el *assignment* en cuestión (que en este caso es "Prueba"), en esta parte vamos a ver un par de íconos relevantes, pero el que nos interesa en este momento es el ícono del birrete que se encuentra sobre la columna *Generate*.

Al presionar este ícono *NBGrader* lleva a cabo ciertas validaciones y en caso de pasarlas se muestra un mensaje indicando que la versión del estudiante se generó de manera correcta, la cual podemos visualizar con el ícono *preview*, como se muestra a continuación.

<img src="./utils/figs/generar.gif"  hspace="5" vspace="5"/>

En caso de tener que modificar la versión del profesor podemos dar click en el ícono *Edit* y eso nos permite realizar las modificaciones necesarias.

Por último es importante notar que la versión del profesor (versión completa) de este *notebook* se encuentra en la ruta *source/Prueba/Evaluacion* y la versión del estudiante se encuentra en la ruta *release/Prueba/Evaluacion*, es importante notar que ambas versiones (profesor y estudiante) de este *notebook* tienen el mismo nombre pero son versiones distintas, por lo que una vez que vayan a ser distribuidos es buena idea renombrarlos para poder identificarlos de manera más sencilla.

<img src="./utils/figs/rutas.gif"  hspace="5" vspace="5"/>

En la siguiente celda puedes ver la evaluación usando assert.

<!-- grade_id: cell-2060978a6c5e20aa, points: 0, grade: false, solution: true -->
```python
def area_cuadrado(lado):
    #BEGIN SOLUTION
    area = lado**2
    #END SOLUTION
    return area
```

<!-- grade_id: cell-251f4b5f896ca1ea, points: 0, grade: true, solution: false -->
```python
"""Verificando el resultado obtenido con la función Trapecio"""
assert area_cuadrado(2) == 4, "Error en el resultado, revisa tu función"

print("¡Correcto, el area del cuadrado es: {}!".format(area_cuadrado(2)))
```

Output:
```
¡Correcto, el area del cuadrado es: 4!
```

---
<a name='ej-1'></a>
#### **<font color="DodgerBlue">Ejercicio 1. Evaluación automatizada *NBGrader*.</font>**

<font color="DarkBlue">Siguiendo los paso descritos previamente, genera un *notebook* que contenga al menos una celda que sea evaluada de manera automatica haciendo uso de *NBGrader*.</font>

---

<a name='2'></a>
## Evaluación MACTI

Para un mejor control de la evaluación, en el proyecto MACTI se han agregado algunos elementos que facilitan aún más la tarea de evaluación del conocimiento, a continuación se describen.

<a name='2-1'></a>
### Estructura evaluación MACTI

Lo primero que debemos tener en cuenta para un mejor control en la evaluación de los conocimientos, es una mejor organización y para eso se sugiere que las soluciones a los ejercicios que se tengan dentro de un *notebook* se almacenen en la ruta *utils/data/NombreNotebook* con el nombre *solmn.npy*, donde:

- *NombreNotebook*: corresponde al *notebook* que se este desarrollando.
- *solmn.npy*: corresponde a la solución del ejerciciomn.

Supongamos que tenemos el tema de "Geometria" (sin acento ya que el nombre de un folder) y dentro de ese tema un *notebook* llamado "FigurasGeometricas" mismo que tiene 2 ejercicios, por lo que cada uno de estos ejercicios debe tener su respectiva solución, de tal forma que el contenido del folder */utils/data/FigurasGeometricas* serían los siguientes elementos.

<img src="./utils/figs/figuras.png"  hspace="5" vspace="5"/>


<a name='2-2'></a>
### Biblioteca *macti_lib*

El siguiente elemento es una biblioteca llamada *macti_lib* desarrollada dentro del marco del proyecto MACTI la cual ya se encuentra instalada en el servidor de MACTI y unicamente tenemos que importarla y hacer uso de ella, sin embargo vale la pena mencionar los siguientes detalles.

**IMPORTANTE:** las soluciones de nuestros ejercicios (*sol01.npy*) son archivos con formato *numpy* que contienen la solución de algún ejercicio en forma de array.

Imaginemos que el ejercicio 1 de nuestro *notebook* "FigurasGeometricas" tiene como solución el número 4, entonces la solución debe ser generada con el método `np.save()` de la biblioteca *numpy* con el nombre *sol01.npy* y almacenado en su respectiva ruta de la siguiente manera.

```python
import numpy as np
# el metodo save cuenta con muchas opciones
np.save("sol01", np.array([4]))
```

Una vez que se generó el archivo *sol01.npy* en el folder donde se está ejecutando el *notebook* en cuestión, solo resta moverlo a la ruta adecuada, como se puede ver a continuación.

<img src="./utils/figs/figurasgeo.gif"  hspace="5" vspace="5"/>

<a name='2-3'></a>
### Método `verifica()`

Ya con la solución en la ruta correcta, solo resta generar un objeto de la clase `Evalua` y utilizar el método `verifica()` de la siguiente forma.

Esta es una prueba para evaluación con **nbgrader**.

Completa la función `area_cuadrado(lado)`, misma que va a recibir la longitu del lado como parámetro y debe devolver el área del cuadrado.

<!-- grade_id: cell-c3762c8877ba0750, points: 0, grade: true, solution: false -->
```python
# importamos la biblioteca
from macti.evaluacion import Evalua
# pasamos la ruta donde se encuentran las soluciones del tema
evaluacion = Evalua('./utils/data/FigurasGeometricas', local=True)
```

Ahora supongamos que en el ejercicio 1 se le pide al alumno definir la función para calcular el área del cuadrado (como vimos previamente).

<!-- grade_id: cell-2060978a6c5e20aa, points: 0, grade: false, solution: true -->
```python
def area_cuadrado(lado):
    #BEGIN SOLUTION
    area = lado**2
    #END SOLUTION
    return area
```

<!-- grade_id: cell-251f4b5f896ca1ea, points: 0, grade: true, solution: false -->
```python
"""
verifica recibe el resultado a verificar (area_cuadrado(2))
y el numero de la solucion, en este caso 1 que corresponde
a la solucion sol01.pny
"""
evaluacion.verifica(np.array([area_cuadrado(2)]), 1)
```

Output:
```
[32m¡Tu resultado es correcto!
```

En caso de que la respuesta (solución del alumno) sea incorrecta, la función `verifica()` muestra una sugerencia de donde se encuentra el error.

<!-- grade_id: cell-2060978a6c5e20aa, points: 0, grade: false, solution: true -->
```python
def area_cuadrado(lado):
    #BEGIN SOLUTION
    area = lado**3
    #END SOLUTION
    return area
```

<!-- grade_id: cell-251f4b5f896ca1ea, points: 0, grade: true, solution: false -->
```python
"""
verifica recibe el resultado a verificar (area_cuadrado(2))
y el numero de la solucion, en este caso 1 que corresponde
a la solucion sol01.pny
"""
evaluacion.verifica(np.array([area_cuadrado(2)]), 1)
```

Output:
```
[39m--------------------------------------------------------------------------------
[31mCuidado: Ocurrió un error en tus cálculos: 
 Lists differ: [8] != [4]

First differing element 0:
8
4

- [8]
+ [4]
[39m--------------------------------------------------------------------------------
```

**IMPORTANTE:** además de escalares (como en el ejemplo anterior) la función verifica puede manejar listas (vectores) o arreglos bidimensionales (matrices).

---
<a name='ej-2'></a>
#### **<font color="DodgerBlue">Ejercicio 2. Evaluación automatizada *MACTI*.</font>**

<font color="DarkBlue">Siguiendo los paso descritos previamente, genera un *notebook* que contenga al menos una celda que sea evaluada de manera automatica haciendo uso de *MACTI*.</font>

---